```python
from random import randint
NROWS = randint(2, 10)
I = randint(0, NROWS-1)
T = [randint(0,100) for _ in range(NROWS)]
val = T[I]
```

:::{admonition} Consigne

Dans le code suivant, quelle valeur devrait avoir `i` pour que `a` prenne la valeur ```{eval}` val ` ``` ?

```
T = {eval}` T `
a = T[i]
```

:::

<!-- grade_id: init, points: 0, grade: false, solution: true -->
```python

```

```python
answer = _
solution = I
answer
```

<!-- grade_id: check, points: 1, grade: true, solution: false -->
```python
assert T[answer] == T[solution]
```