```python
from random import randint
NROWS = randint(2, 5)
NCOLUMNS = randint(2, 5)
I1 = randint(0, NROWS-1)
I2 = randint(0, NCOLUMNS-1)
T = [[randint(0,100) for _ in range(NCOLUMNS)] for _ in range(NROWS)]
```

:::{admonition} Consigne

Observez le code suivant :
```
T = {eval}` T `
a = T[{eval}` I1 `][{eval}` I2 `]
```

Quelle est la valeur de `a` ?

:::

<!-- grade_id: init, points: 0, grade: false, solution: true -->
```python

```

```python
answer = _
solution = T[I1][I2]
answer
```

<!-- grade_id: check, points: 1, grade: true, solution: false -->
```python
assert answer == solution
```