```python
from jupylates.jupylates_helpers import SUBSTITUTE
import random

var = random.choice("abcijlkpq")
varT = random.choice("st")

seuil = random.randint(3, 5)
condition = f"{var} {random.choice(['<', '>'])} {seuil}"
expression = random.choice([var, f"{var}**2", f"2*{var}"])
COMPREHENSION = f"[{expression} for {var} in {varT} if {condition}]"

SUBSTITUTE(
    T=varT,
    COMPREHENSION=COMPREHENSION)
```

```python
T = [random.randint(0, 10)
     for _ in range(random.randint(3, 5))]
```

:::{admonition} Consigne

Soit `T` le tableau {eval}`T`. Quelle est la valeur de l'expression suivante?

``` python
COMPREHENSION
```

:::

<!-- grade_id: init, points: 0, grade: false, solution: true -->
```python

```

```python
answer = _
```

<!-- grade_id: check, points: 1, grade: true, solution: false -->
```python
assert answer == COMPREHENSION
```