```python
from random import randint
NROWS = randint(2, 10)
I = randint(0, NROWS-1)
T = [randint(0,100) for _ in range(NROWS)]
```

:::{admonition} Consigne

Observez le code suivant :
```
T = {eval}` T `
a = T[{eval}` I `]
```

Quelle est la valeur de `a` ?

:::

<!-- grade_id: init, points: 0, grade: false, solution: true -->
```python

```

```python
answer = _
solution = T[I]
answer
```

<!-- grade_id: check, points: 1, grade: true, solution: false -->
```python
assert answer == solution
```