```python
from random import randint
NROWS = randint(2, 4)
T = [randint(0,100) for _ in range(NROWS)]
i = randint(0, NROWS-1)

to_append = randint(0,100)
app = "a = " + str(T) + "\na.append(" + str(to_append) + ")"
solution_app = T.copy()
solution_app.append(to_append)


to_remove = T[i]
rem = "a = " + str(T) + "\na.remove(" + str(to_remove) + ")"
solution_rem = T.copy()
solution_rem.remove(to_remove)


to_insert = randint(0,100)
ins = "a = " + str(T) + "\na.insert(" + str(i) +", " + str(to_insert) + ")"
solution_ins = T.copy()
solution_ins.insert(i, to_insert)


valeur = T[i]
ind = "T = " + str(T) + "\na = T.index(" + str(valeur)+ ")"
solution_ind = T.index(valeur)

de = "a = " + str(T) + "\ndel a[" + str(i)+ "]"
solution_de = T.copy()
del solution_de[i]

le = "T = " + str(T) + "\na = len(T)"
solution_le = len(T)

maxi = "T = " + str(T) + "\na = max(T)"
solution_maxi = max(T)

mini = "T = " + str(T) + "\na = min(T)"
solution_mini = min(T)

variantes = [app, rem, ins, ind, de, le, maxi, mini]
solutions = [solution_app, solution_rem, solution_ins, solution_ind, solution_de, solution_le,
            solution_maxi, solution_mini]
exercice = randint(0,len(variantes)-1)
```

:::{admonition} Consigne

Quel est le contenu de `a` après l'exécution du code ci-dessous ?


```
{eval}` variantes[exercice] `
```

:::

<!-- grade_id: init, points: 0, grade: false, solution: true -->
```python

```

```python
answer = _
solution = solutions[exercice]
answer
```

<!-- grade_id: check, points: 1, grade: true, solution: false -->
```python
assert answer == solution
```