```python
from random import randint
NROWS1 = randint(1, 4)
NROWS2 = randint(1, 4)
T1 = [randint(0,100) for _ in range(NROWS1)]
T2 = [randint(0,100) for _ in range(NROWS2)]


add = "l1 = " + str(T1) + "\nl2 = " + str(T2) + "\na = l1 + l2"
solution_add = T1 + T2

m = randint(1, 3)
mul = "l = " + str(T1) + "\na = l * " + str(m)
solution_mul = T1 * m

ext = "a = " + str(T1) + "\nb = " + str(T2) + "\na.extend(b)"
solution_ext = T1.copy()
solution_ext.extend(T2)

variantes = [add, mul, ext]
solutions = [solution_add, solution_mul, solution_ext]
exercice = randint(0,len(variantes)-1)
```

:::{admonition} Consigne

Quel est le contenu de `a` après l'exécution du code ci-dessous ?


```
{eval}` variantes[exercice] `
```

:::

<!-- grade_id: init, points: 0, grade: false, solution: true -->
```python

```

```python
answer = _
solution = solutions[exercice]
answer
```

<!-- grade_id: check, points: 1, grade: true, solution: false -->
```python
assert answer == solution
```