:::{admonition} Objectifs pédagogiques
Utiliser la boucle for avec compteur et accumulateur.
:::

```python
from random import randint
import math
```

::::{admonition} Consigne

Définissez une fonction "puissance" qui prend en paramètre des entiers
positifs `a` et `b` et calcule `a^b`.

:::{tip} Indication
:class: dropdown

Utilisez une boucle `for` avec un compteur et un accumulateur. 

% Voir le cours Boucles/03-compteurs-accumulateurs.md

:::

::::

<!-- grade_id: init, points: 0, grade: false, solution: true -->
```python
def puissance(a, b):
    résultat = 1

    for i in range(b):
        résultat *= a   # Équivalent à résultat = résultat * a

    return résultat
```

<!-- grade_id: check, points: 1, grade: true, solution: false -->
```python
assert puissance(2, 3) == 2**3, f"Valeur attendue : {2**3}, valeur obtenue : {puissance(2, 3)}"
assert puissance(-3, 2) == (-3)**2, f"Valeur attendue : {(-3)**2}, valeur obtenue : {puissance(-3, 2)}"
```