:::{admonition} Objectifs pédagogiques
Utiliser les conditionnelles
:::

```python
from random import randint
import math
from jupylates.jupylates_helpers import RANDOM_INT, CONST, INPUT
```

:::{admonition} Consigne

Définir une fonction `est_triangle_rectangle(a, b, c)` prenant en
entrée trois entiers `a`, `b` et `c` représentant les longueurs des
côtés d'un triangle, et renvoyant `True` si le triangle est rectangle
d'hypothénuse `c` et `False` sinon.

**Indication :** vérifier que $c^2 == a^2 + b^2$.

:::

<!-- grade_id: init, points: 0, grade: false, solution: true -->
```python
def est_triangle_rectangle(a, b, c):
    # Vérifier le théorème de Pythagore
    return c ** 2 == a ** 2 + b ** 2
```

<!-- grade_id: check, points: 1, grade: true, solution: false -->
```python
assert est_triangle_rectangle(3, 4, 5), f"est_triangle_rectangle(3, 4, 5) devrait renvoyer True"
#assert est_triangle_rectangle(5, 4, 3), f"est_triangle_rectangle(5, 4, 3) devrait renvoyer True"
assert est_triangle_rectangle(6, 8, 10), f"est_triangle_rectangle(6, 8, 10) devrait renvoyer True"
#assert est_triangle_rectangle(15, 12, 20), f"est_triangle_rectangle(15, 12, 20) devrait renvoyer True"
assert not est_triangle_rectangle(2, 4, 5), f"est_triangle_rectangle(2, 4, 5) devrait renvoyer False"
assert not est_triangle_rectangle(200, 16, 12), f"est_triangle_rectangle(200, 16, 12) devrait renvoyer False"
```