:::{admonition} Objectifs pédagogiques
Utiliser les conditionnelles
:::

```python
from random import randint
import math
from jupylates.jupylates_helpers import RANDOM_INT, CONST, INPUT
```

:::{admonition} Consigne

Définir une fonction `est_triangle_rectangle(a, b, c)` prenant en
entrée trois entiers `a`, `b` et `c` qui représentent les longueurs
des côtés d'un triangle, et renvoyant `True` si le triangle est
rectangle, `False` sinon.

**Indication :** Cette fois, l'hypoténuse peut être en `a`, `b` ou
`c`.

:::

<!-- grade_id: init, points: 0, grade: false, solution: true -->
```python
def est_triangle_rectangle(a, b, c):
    # Vérifier le théorème de Pythagore dans les trois cas
    return a**2 == b**2 + c**2 or \
           b**2 == a**2 + c**2 or \
           c**2 == a**2 + b**2
```

<!-- grade_id: check, points: 1, grade: true, solution: false -->
```python
assert est_triangle_rectangle(8, 10, 6), f"est_triangle_rectangle(8, 10, 6) devrait renvoyer True"
assert not est_triangle_rectangle(10, 7, 6), f"est_triangle_rectangle(10, 7, 6) devrait renvoyer False"
assert est_triangle_rectangle(16, 20, 12), f"est_triangle_rectangle(16, 20, 12) devrait renvoyer True"
assert not est_triangle_rectangle(200, 16, 12), f"est_triangle_rectangle(200, 16, 12) devrait renvoyer False"
```