:::{admonition} Objectifs pédagogiques
échange de variables
:::

```python
from jupylates.jupylates_helpers import RANDOM_INT, CONST, INPUT
I1: CONST = RANDOM_INT(0, 7)
I2: CONST = RANDOM_INT(8, 15)
```

```python
X = I1
Y = I2

Z = X
X = Y
Y = Z
r = Y
```

:::{admonition} Consigne

Quelle est la valeur attendue de r?

:::

<!-- grade_id: init, points: 0, grade: false, solution: true -->
```python
result = INPUT(
    ### BEGIN SOLUTION
    Y
    ### END SOLUTION
)
```

<!-- grade_id: check, points: 1, grade: true, solution: false -->
```python
assert result == r
```